# Numpy Exercise 1

## Imports

```python
import numpy as np
%matplotlib inline
import matplotlib.pyplot as plt
import seaborn as sns
```

```python
import antipackage
import github.ellisonbg.misc.vizarray as va
```

## Checkerboard

Write a Python function that creates a square `(size,size)` 2d Numpy array with the values `0.0` and `1.0`:

* Your function should work for both odd and even `size`.
* The `0,0` element should be `1.0`.
* The `dtype` should be `float`.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
c = np.arange(0, 10, 2, dtype=np.float)
c


data = [1.0,0.0],[0.0,1.0]
a = np.array(data)
va.vizarray(a)
```

Output:
```
<ipythonblocks.ipythonblocks.BlockGrid at 0x7fbbe8323c10>
```

```python
def checkerboard(x,y):
    a = np.ones((x,y))
    while x<2:
        
```

Output:
```
<ipythonblocks.ipythonblocks.BlockGrid at 0x7fbc01b97b50>
```

<!-- grade_id: numpyex01a, points: 6, grade: true, solution: false -->
```python
a = checkerboard(4)
assert a[0,0]==1.0
assert a.sum()==8.0
assert a.dtype==np.dtype(float)
assert np.all(a[0,0:5:2]==1.0)
assert np.all(a[1,0:5:2]==0.0)

b = checkerboard(5)
assert b[0,0]==1.0
assert b.sum()==13.0
assert np.all(b.ravel()[0:26:2]==1.0)
assert np.all(b.ravel()[1:25:2]==0.0)
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-7-4aa63174c227> in <module>()
      1 a = checkerboard(4)
----> 2 assert a[0,0]==1.0
      3 assert a.sum()==8.0
      4 assert a.dtype==np.dtype(float)
      5 assert np.all(a[0,0:5:2]==1.0)

TypeError: 'NoneType' object has no attribute '__getitem__'
```

Use `vizarray` to visualize a checkerboard of `size=20` with a block size of `10px`.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: numpyex01b, points: 2, grade: true, solution: false -->
```python
assert True
```

Use `vizarray` to visualize a checkerboard of `size=27` with a block size of `5px`.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: numpyex01c, points: 2, grade: true, solution: false -->
```python
assert True
```